Before you submit this tick, make sure everything runs as expected. First, **save your work** (in the menubar, select File$\rightarrow$Save and Checkpoint). Then **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", and remove any instances of `failwith "Not implemented";;`


---

# FoCS Tick 1* - Recursive Functions Continued

Write an OCaml function `eapprox(n)` to sum the `n` terms in the approximation
$$\exp \approx 1 + \frac{1}{1!} + \frac{1}{2!} + ... + \frac{1}{(n-1)!}$$
Again, each term can be cheaply computed.


<!-- grade_id: cell-ad5fdb2dd501d4fa, points: 0, grade: false, solution: true -->
```OCaml
let eapprox n =
    let rec hlp n x count = match count with
    | 1 -> x
    | _ -> x +. hlp (n+1) (x /. (float_of_int n)) (count-1)
    in hlp 1 1.0 n
```

Output:
```
val eapprox : int -> float = <fun>
```

Test your function with the following cell.

<!-- grade_id: cell-be54a6efe6d5abb1, points: 1, grade: true, solution: false -->
```OCaml
let approx_eq a b = abs_float (a -. b) < 1e-10;;

assert (approx_eq (eapprox 1) 1.0);;
assert (approx_eq (eapprox 2) 2.0);;
assert (approx_eq (eapprox 3) 2.5);;
assert (approx_eq (eapprox 10) 2.7182815255);;
```

Output:
```
val approx_eq : float -> float -> bool = <fun>

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()
```

2. Write an OCaml function `exp z n` to compute the exponential function:
$$\exp(z) \approx  1 + \frac{z}{1!} + \frac{z^2}{2!} + ... + \frac{z^{n-1}}{(n-1)!}$$

<!-- grade_id: cell-5874508d2d04bbb2, points: 0, grade: false, solution: true -->
```OCaml
let exp z n =
    let rec hlp z n x count = match count with
    | 1 -> x
    | _ -> x +. hlp z (n+1) (x /. (float_of_int n) *. z) (count-1)
    in hlp z 1 1.0 n
```

Output:
```
val exp : float -> int -> float = <fun>
```

Check your solution with the following test cell.

<!-- grade_id: cell-b09a75ce5244fb87, points: 1, grade: true, solution: false -->
```OCaml
assert (approx_eq (exp 1.0 1) 1.0);;
assert (approx_eq (exp 1.0 2) 2.0);;
assert (approx_eq (exp 1.0 3) 2.5);;
assert (approx_eq (exp 1.0 10) 2.7182815255);;
assert (approx_eq (exp 1.5 18) (Stdlib.exp 1.5))
```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()
```

The *binary greatest common divisor algorithm* is based on the following observations:
$$\begin{align}
gcd(a,a) &= a \\
gcd(2a,2b) &= 2 gcd(a,b) \\
gcd(2a+1,2b) &= gcd(2a+1,b) \\
gcd(2a+1,2b+1) &= gcd(2b+1,a-b)
\end{align}$$
This algorithm is particularly appealing to implement in hardware due to its reliance on division by two.

3. Fill in the gaps in this description and implement the resulting recursive function `gcd` in OCaml, to work with positive integers. You will need to ensure that zero and negative integers are never introduced in recursive calls.

<!-- grade_id: cell-e60d499a2aaee53f, points: 0, grade: false, solution: true -->
```OCaml
let even n = ((n mod 2) = 0)

let rec gcd a b =
    if b>a then gcd b a
    else if a=b then a
    else if even a && even b then 2 * gcd (a/2) (b/2)
    else if (not (even a)) && even b then gcd a (b/2)
    else if even a && not (even b) then gcd (a/2) b
    else if (not (even a)) && (not (even b)) then gcd b ((a-1)/2 - (b-1)/2)
    else 0

```

Test your implementation with the following cell.

<!-- grade_id: cell-6f10c05615dac8ac, points: 1, grade: true, solution: false -->
```OCaml
assert(gcd 1 10 = 1);;
assert(gcd 2 10 = 2);;
assert(gcd 3 10 = 1);;
assert(gcd 1 9 = 1);;
assert(gcd 2 9 = 1);;
assert(gcd 3 9 = 3);;
assert(gcd 24826148 45296490 = 526);;
assert(gcd 461952 116298 = 18);;
```

```OCaml

```

```OCaml

```