Before you submit this tick, make sure everything runs as expected. First, **save your work** (in the menubar, select File$\rightarrow$Save and Checkpoint). Then **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", and remove any instances of `failwith "Not implemented";;`


---

# FoCS Tick 2* - Lists of Lists

Write a function `choose k xs` that returns all k-element lists that can be drawn from xs, ignoring the order of list elements. If `n` is the length of `xs`, then (provided $k\le n$) the result should be an $\begin{pmatrix} n \\ k \end{pmatrix}$-element list.

It may be useful to first define two auxiliary functions: one that adds a specific element to every list in a list of lists:

<!-- grade_id: cell-5063e744d0696f60, points: 0, grade: false, solution: true -->
```OCaml
let allcons x xs = List.map (fun l -> x::l) xs
```

Output:
```
val allcons : 'a -> 'a list list -> 'a list list = <fun>
```

<!-- grade_id: cell-6641ad5ec55c6bc4, points: 1, grade: true, solution: false -->
```OCaml
assert(allcons 1 [[]; [2]; [2;3]] = [[1]; [1;2]; [1;2;3]]);;
```

Output:
```
- : unit = ()
```

and a function that concatenates two lists (of lists) together:

<!-- grade_id: cell-830e942cb0cd494d, points: 0, grade: false, solution: true -->
```OCaml
let rec concat xs ys = xs @ ys
```

Output:
```
val concat : 'a list -> 'a list -> 'a list = <fun>
```

<!-- grade_id: cell-89ab2833a874354c, points: 1, grade: true, solution: false -->
```OCaml
assert(concat [[1];[2;3]] [[];[4;5;6]] = [[1];[2;3];[];[4;5;6]]);;
```

Output:
```
- : unit = ()
```

although it is not mandatory to use these functions in your solution.

<!-- grade_id: cell-c89f9b4b13e96987, points: 0, grade: false, solution: true -->
```OCaml
let rec choose k xs = if k=(List.length xs) then [xs] else match xs with
    | [] -> []
    | x::xs -> allcons x (choose k xs)

```

Output:
```
val choose : int -> 'a list -> 'a list list = <fun>
```

The following cell contains an inefficient test for the choose function. Once again you are not expected to understand it in detail.

<!-- grade_id: cell-2e357dfbc562cd1b, points: 1, grade: true, solution: false -->
```OCaml
let compare is js =
    let rec check ys = function
      | x::xs ->
          let sorted_x = List.sort Stdlib.compare x in
          let ok = List.exists (fun y ->
              let sorted_y = List.sort Stdlib.compare y in
              sorted_x = sorted_y) ys in
          ok && check ys xs
      | _ -> true
    in
    check is js && check js is && (List.length is = List.length js);;

let result = choose 4 [1;2;3;4;5;6];;
let correct = [[3; 4; 5; 6]; [2; 4; 5; 6]; [2; 3; 5; 6]; [2; 3; 4; 6]; [2; 3; 4; 5];
   [1; 4; 5; 6]; [1; 3; 5; 6]; [1; 3; 4; 6]; [1; 3; 4; 5]; [1; 2; 5; 6];
   [1; 2; 4; 6]; [1; 2; 4; 5]; [1; 2; 3; 6]; [1; 2; 3; 5]; [1; 2; 3; 4]];;
assert(compare result correct)

```

Output:
```
val compare : 'a list list -> 'a list list -> bool = <fun>

val result : int list list = [[1; 2; 3; 4; 5; 6]]

val correct : int list list =
  [[3; 4; 5; 6]; [2; 4; 5; 6]; [2; 3; 5; 6]; [2; 3; 4; 6]; [2; 3; 4; 5];
   [1; 4; 5; 6]; [1; 3; 5; 6]; [1; 3; 4; 6]; [1; 3; 4; 5]; [1; 2; 5; 6];
   [1; 2; 4; 6]; [1; 2; 4; 5]; [1; 2; 3; 6]; [1; 2; 3; 5]; [1; 2; 3; 4]]

Exception: Assert_failure ("[8]", 17, 0).
Called from Stdlib__Fun.protect in file "fun.ml", line 33, characters 8-15
Re-raised at Stdlib__Fun.protect in file "fun.ml", line 38, characters 6-52
Called from Topeval.load_lambda in file "toplevel/byte/topeval.ml", line 89, characters 4-150
```

```OCaml

```