Before you submit this tick, make sure everything runs as expected. First, **save your work** (in the menubar, select File$\rightarrow$Save and Checkpoint). Then **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", and remove any instances of `failwith "Not implemented";;`


---

# FoCS Tick 1 - Recursive Functions

OCaml is a strongly typed language and distinguishes between floating point and integer values. The arithmetic operators `+`, `-`, `*` and `/` operate on integers, whereas the equivalent floating point operators are `+.`, `-.`, `*.` and `/.`. OCaml will interpret numbers as integers unless there is a decimal point or exponent in them, so the following are integers: `0`, `1`, `1000`, and these are floating point numbers: `1.0`, `1.`, `1e0`. The functions `int_of_float` and `float_of_int` can be used to convert between the types.

Everybody who writes computer programs should be aware of floating-point rounding errors. Ask OCaml to evaluate the following expression, and carefully note its response.

```OCaml
1.0 -. 0.9 -. 0.1
```

Output:
```
- : float = -2.77555756156289135e-17
```

The two solutions of the equation $ax^2 + bx + c = 0$ are given by the quadratic formula,
$$\frac{−b \pm \sqrt{b^2 −4ac}}{2a}$$
The following OCaml code computes one of these roots.

```OCaml
let rootplus a b c = (-. b +. sqrt (b *. b -. 4.0 *. a *. c)) /. (2.0 *. a)
```

Output:
```
val rootplus : float -> float -> float -> float = <fun>
```

1. Write an OCaml function `evalquad` such that `evalquad a b c x` computes the value of $ax^2+bx+c$.

<!-- grade_id: cell-6971eca02b2d6bb0, points: 0, grade: false, solution: true -->
```OCaml
let evalquad a b c x =
    a*.x*.x +. b*.x +. c
```

Output:
```
val evalquad : float -> float -> float -> float -> float = <fun>
```

The following cell contains tests for your evalquad function so you can check your own solution before submitting it. Don't worry if you don't understand the cell - the important this is that when the cell is executed it will display an error if the tests don't pass. 

<!-- grade_id: cell-07fbcdd398bee7bc, points: 1, grade: true, solution: false -->
```OCaml
let approx_eq a b = abs_float (a -. b) < 1e-10;;
assert(approx_eq (evalquad 1.2 3.4 5.6 7.8) 105.128)
```

Output:
```
val approx_eq : float -> float -> bool = <fun>

- : unit = ()
```

Mathematical function definitions are often recursive. The well-known factorial function, $n!$, is defined by $0!=1$ and (for n > 0) $$n! = n * (n-1)!$$

2. Write an OCaml function `facr` to compute the integer $n!$ by recursion, and also a function `faci` to compute the integer $n!$ by iteration (in the sense described in Lecture 2). When writing your function, consider what to do with negative arguments.

<!-- grade_id: cell-8f0893754b7065bb, points: 0, grade: false, solution: true -->
```OCaml
exception NegativeFactorial of int

let rec facr n = 
    if n<0 then raise (NegativeFactorial n)
    else
        match n with
        | 0 -> 1
        | 1 -> 1
        | _ -> n * facr (n-1)

let faci n =
    if n<0 then raise (NegativeFactorial n)
    else
        let rec hlp n acc = match n with
            | 0 -> 1
            | 1 -> acc
            | _ -> hlp (n-1) (n*acc)
        in hlp n 1
    
```

Output:
```
exception NegativeFactorial of int

val facr : int -> int = <fun>

val faci : int -> int = <fun>
```

The following cell contains tests for `facr` and `faci`.

<!-- grade_id: cell-7173a13d67176a6b, points: 1, grade: true, solution: false -->
```OCaml
assert (facr 10 = 3628800);;
assert (faci 10 = 3628800);;
assert (facr 11 = 39916800);;
assert (faci 11 = 39916800);;

```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()
```

3. Write an ML function `sumt` to sum the first n terms of the geometric series
$$1+\frac{1}{2}+\frac{1}{4} + ... + \frac{1}{2^{n-1}} + ...$$
for $n>0$. For example, when $n=2$ the sum is $\frac{1}{2^0}+\frac{1}{2^1}$, or 1.5
Observe that each term can be cheaply computed from its predecessor. A mathematician might remark that it is useful to consider the slightly more general function
$$f(x,n) = x + \frac{x}{2} + \frac{x}{4} + ... + \frac{x}{2^{n-1}}$$
This function satisfies the recursive definition (for $n>0$)
$$f(x,n) = x + f(\frac{x}{2}, n-1)$$

<!-- grade_id: cell-7e42f38971ae316a, points: 0, grade: false, solution: true -->
```OCaml
let sumt n = 
    let rec hlp x n = match n with
        | 1 -> x
        | _ -> x +. hlp (x /. 2.0) (n-1)
    in hlp 1.0 n
    
    
```

Output:
```
val sumt : int -> float = <fun>
```

Test your function with the following cell.

<!-- grade_id: cell-25e0ad2a629e22e0, points: 1, grade: true, solution: false -->
```OCaml
assert (approx_eq (sumt 2) (1.5));;
assert (approx_eq (sumt 3) (1.75));;
assert (approx_eq (sumt 4) (1.875));;
```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()
```

```OCaml

```