Before you submit this tick, make sure everything runs as expected. First, **save your work** (in the menubar, select File$\rightarrow$Save and Checkpoint). Then **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", and remove any instances of `failwith "Not implemented";;`


---

# FoCS Tick 4 - Functions as Data, Integer Streams

In OCaml, functions can be given as input and can be returned as results from functions. As coded below, the function twice takes the argument f and returns an anonymous function that applies f twice to its argument.

```OCaml
let twice f = fun x -> f (f x)
```

Output:
```
val twice : ('a -> 'a) -> 'a -> 'a = <fun>
```

Can you see why the following expression evaluates to `11`?

```OCaml
twice (twice (twice (fun i -> i+1))) 3;;
```

If $f$ is a function and $n \ge 0$ is an integer then the function $f^n$ is defined as follows:
$$f^n(x) = \underbrace{f(f(...f(}_{n\text{ times}}x)...))$$
In particular, $f^0(x)=x$. Given that $s$ is the function such that $s(x) = x + 1$ (which adds 1 to its argument), we can express the sum of two non-negative integers $m$ and $n$ as $m+n=s^n(m)$ (i.e. 1 is added to $m$ but $n$ times over).

Think about how to express the product $m\times n$ and power $m^n$ similarly, for non-negative integers. $Hint:$ Consider what has to be repeated $n$ times over to obtain $m \times n$ or $m^n$. Note that the functions that are analogous to $s(x)$ may have to depend upon $m$.

1. Write a curried function `nfold` such that `nfold f n` returns the function $f^n$. Use `nfold` to write curried functions `sum`, `product` and `power` to compute integer sums, products and powers as outlined above.

<!-- grade_id: cell-df350e67125dbd51, points: 0, grade: false, solution: true -->
```OCaml
let rec nfold f n x = match n with 
 | 0 -> x
 | n -> nfold f (n-1) (f x)
    
let sum m n = nfold (fun a -> a+1) n m
    
let product m n = nfold (sum m) n 0

let power m n = nfold (product m) n 1
```

Output:
```
val nfold : ('a -> 'a) -> int -> 'a -> 'a = <fun>

val sum : int -> int -> int = <fun>

val product : int -> int -> int = <fun>

val power : int -> int -> int = <fun>
```

<!-- grade_id: cell-ffd0efb8a9b27ea9, points: 1, grade: true, solution: false -->
```OCaml
assert(sum 0 0 = 0);;
assert(sum 1 0 = 1);;
assert(sum 0 1 = 1);;
assert(sum 1 1 = 2);;
assert(sum 100 100 = 200);;
assert(product 0 0 = 0);;
assert(product 1 0 = 0);;
assert(product 0 1 = 0);;
assert(product 1 1 = 1);;
assert(product 14 9 = 126);;
assert(power 0 1 = 0);;
assert(power 1 0 = 1);;
assert(power 1 1 = 1);;
assert(power 3 0 = 1);;
assert(power 3 5 = 243);;
```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()
```

Here is a definition of streams (a.k.a sequences, infinite or lazy lists) that cannot terminate. Note that the function `from` (which creates the stream of integers starting from a specified value) can be declared exactly the same as with the type `'a seq` presented in the lectures.

```OCaml
type 'a stream = Cons of 'a * (unit -> 'a stream)
let rec from k = Cons(k, fun () -> from (k+1))
```

Output:
```
type 'a stream = Cons of 'a * (unit -> 'a stream)

val from : int -> int stream = <fun>
```

2. Write a function `nth s n` to return the nth element of `s`. For example, `nth (from 1) 100` should return 100.

<!-- grade_id: cell-24a79eeb16fa121e, points: 0, grade: false, solution: true -->
```OCaml
let rec nth s n = match s, n with
    | Cons(x, xf), 1 -> x
    | Cons(x, xf), n -> nth (xf ()) (n-1)
```

Output:
```
val nth : 'a stream -> int -> 'a = <fun>
```

<!-- grade_id: cell-72978811a2d2d59c, points: 1, grade: true, solution: false -->
```OCaml
assert(nth (from 1) 100 = 100)
```

Output:
```
- : unit = ()
```

3. Make the stream `squares` of positive squares (1, 4, 9, ...) and find its 49th element.

<!-- grade_id: cell-6139ccc02d8db255, points: 0, grade: false, solution: true -->
```OCaml
let rec mapseq f = function
    | Cons(x, xf) -> Cons (f x, fun () -> mapseq f (xf()))

let squares = mapseq (fun a -> a*a) (from 1)
```

Output:
```
val mapseq : ('a -> 'b) -> 'a stream -> 'b stream = <fun>

val squares : int stream = Cons (1, <fun>)
```

<!-- grade_id: cell-0c4b85b121c5a03e, points: 1, grade: true, solution: false -->
```OCaml
assert(nth squares 49 = 2401);;
```

Output:
```
- : unit = ()
```

4. Write a function `map2 f xs ys`, similar to `mapq`, to take $x_1, x_2, x_3,...$ and $y_1, y_2, y_3, ...$ and return the stream $f(x_1)(y_1), f(x_2)(y_2), f(x_3)(y_3), ...$

<!-- grade_id: cell-052ee5bc39383529, points: 0, grade: false, solution: true -->
```OCaml
let rec map2 f xs ys = match xs, ys with
    | Cons(x, xf), Cons(y, yf) -> Cons(f x y, fun () -> map2 f (xf()) (yf()) )

```

Output:
```
val map2 : ('a -> 'b -> 'c) -> 'a stream -> 'b stream -> 'c stream = <fun>
```

<!-- grade_id: cell-fbe5af6c3285b733, points: 1, grade: true, solution: false -->
```OCaml
let rec take n xs =
    match xs with
    | Cons (x,f) ->
        if n=0
        then []
        else x::(take (n-1) (f()));;
assert(take 5 (map2 (fun x y -> x+y) squares (from 10)) = [ 11; 15; 21; 29; 39]);;
```

Output:
```
val take : int -> 'a stream -> 'a list = <fun>

- : unit = ()
```

```OCaml

```