Before you submit this tick, make sure everything runs as expected. First, **save your work** (in the menubar, select File$\rightarrow$Save and Checkpoint). Then **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", and remove any instances of `failwith "Not implemented";;`


---

# FoCS Tick 2 - Structured Data - Pairs and Lists

Before tackling the questions, try some simple experiments with structured data. Evaluate the following selector functions:

```OCaml
let fst (x, y) = x
let snd (x, y) = y
```

Output:
```
val fst : 'a * 'b -> 'a = <fun>

val snd : 'a * 'b -> 'b = <fun>
```

To experiment with them, evaluate the following value declarations:

```OCaml
let p = ("red", 3)
let q = (p, "blue")
let r = (q, p)
let s = ((23, "grey"), r)
```

Output:
```
val p : string * int = ("red", 3)

val q : (string * int) * string = (("red", 3), "blue")

val r : ((string * int) * string) * (string * int) =
  ((("red", 3), "blue"), ("red", 3))

val s : (int * string) * (((string * int) * string) * (string * int)) =
  ((23, "grey"), ((("red", 3), "blue"), ("red", 3)))
```

Before evaluating the following cell, predict what the result of each statement will be. Extend it with your own statements.

```OCaml
fst (fst q);;
```

Output:
```
- : string = "red"
```

```OCaml
fst (fst p);;
```

Output:
```
File "[4]", line 1, characters 4-11:
1 | fst (fst p);;
        ^^^^^^^
Error: This expression has type string but an expression was expected of type
         'a * 'b
```

```OCaml
fst (snd s);;
```

Output:
```
- : (string * int) * string = (("red", 3), "blue")
```

Note that triples are not pairs! Compare OCaml's response to each of the following:

```OCaml
fst ((1,2),3)
```

Output:
```
- : int * int = (1, 2)
```

```OCaml
fst (1,2,3)
```

Output:
```
File "[7]", line 1, characters 4-11:
1 | fst (1,2,3)
        ^^^^^^^
Error: This expression has type 'a * 'b * 'c
       but an expression was expected of type 'd * 'e
```

Recall that, although lists can hold structured data, all elements in a list must have the same *type*. Note OCaml's response to each of the following. Remember that `1.0` is a `float` and `2` is an `int`

```OCaml
["orange";1;2]
```

Output:
```
File "[8]", line 1, characters 10-11:
1 | ["orange";1;2]
              ^
Error: This expression has type int but an expression was expected of type
         string
```

```OCaml
1.0::[2]
```

Output:
```
File "[9]", line 1, characters 6-7:
1 | 1.0::[2]
          ^
Error: This expression has type int but an expression was expected of type
         float
  Hint: Did you mean `2.'?
```

```OCaml
[3; (3,3)]
```

Output:
```
File "[10]", line 1, characters 4-9:
1 | [3; (3,3)]
        ^^^^^
Error: This expression has type 'a * 'b
       but an expression was expected of type int
```

*Important note*: In the questions below there are inputs that make no sense. For example, asking for the last element of an empty list. The correct way to deal with such problems is to raise an exception. If you already know how to do this, then please do so. If you do not, you may omit these cases. Your functions will produce a warning about matches not being exhaustive but that is acceptable for this tick and you will still pass the automated checker.

1. The function `List.hd` returns the first element of a list. Getting at the last element is harder. Write a recursive function `last` to return the last element of a list. For example, on input `[1;2;3]`, your function `last` should return `3`

<!-- grade_id: cell-63a0a50894cc1683, points: 0, grade: false, solution: true -->
```OCaml
exception EmptyList

let rec last = function
    | x1::x2::xs -> last (x2::xs)
    | x::xs -> x
    | [] -> raise EmptyList
```

Output:
```
exception EmptyList

val last : 'a list -> 'a = <fun>
```

<!-- grade_id: cell-5711db27280c84bf, points: 1, grade: true, solution: false -->
```OCaml
assert(last [1;2;3] = 3);;
assert(last [10;9;8;7;6;5;4;3;2;1001] = 1001);;
assert(last [1] = 1);;
```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()
```

Now do the same thing for `List.tl`, i.e. write a recursive function `butlast` to remove the last element of a list. For example, `butlast [1;2;3;4]` should return `[1;2;3]`. Note that `butlast xs` must be equivalent to `List.rev (List.tl (List.rev xs))`, so `butlast [1]` should return `[]`. Check that you can compute
the time and space complexity of `butlast` (and compare it with `List.rev (List.tl (List.rev xs))`).

<!-- grade_id: cell-0e62a9655ced4fbd, points: 0, grade: false, solution: true -->
```OCaml
let rec butlast = function
    | [] -> []
    | x1::x2::xs -> x1 :: butlast(x2::xs)
    | x -> []
```

Output:
```
val butlast : 'a list -> 'a list = <fun>
```

<!-- grade_id: cell-71b53e441e1348c7, points: 1, grade: true, solution: false -->
```OCaml
assert(butlast [1;2;3;4] = [1;2;3]);;
assert(butlast [1] = []);;
```

Output:
```
- : unit = ()

- : unit = ()
```

3. Write a function `nth` such that `nth xs n` returns the *n*th element of list xs, counting the head of the list as element zero.

<!-- grade_id: cell-cd66c2343ca9ebf7, points: 0, grade: false, solution: true -->
```OCaml
exception ListIndexOutOfBounds

let rec nth xs n = match n, xs with
    | (k, []) -> raise ListIndexOutOfBounds
    | (0, x::xs) -> x
    | (k, x::xs) -> nth xs (k-1)
    
```

Output:
```
exception ListIndexOutOfBounds

val nth : 'a list -> int -> 'a = <fun>
```

<!-- grade_id: cell-709e6f6f27bb69d4, points: 1, grade: true, solution: false -->
```OCaml
assert (nth [1;2;3;4] 0 = 1);;
assert (nth [1;2;3;4] 1 = 2);;
assert (nth [1;2;3;4] 2 = 3);;
assert (nth [1;2;3;4] 3 = 4);;
assert (nth [1;2;3;4;5;6;7;8;9;10] 7 = 8);;
```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()
```

```OCaml

```